# Empezando

Temas:

1. Cómo imprimir
1. Cómo asignar variables
1. Cómo comentar
1. Sintaxis para matemáticas básicas


## Cómo imprimir

En Julia, comúnmente se usa `println()` para imprimir, pero `print()` también funciona. ¿Cuál es la diferencia?

```julia
println("¡Hola, CORE 2020!")
```

Output:
```
¡Hola, CORE 2020!
```

## Cómo asignar variables

Solo se necesita un nombre de variable, un valor y un signo de igual.

Julia se encarga de los *tipos* por nosotros.

```julia
my_answer = 42
typeof(my_answer)
```

Output:
```
Int64
```

```julia
my_pi = 3.14159
typeof(my_pi)
```

Output:
```
Float64
```

**Pendiente de ver variables *Unicode***

Después de asignar un valor a una variable, se puede reasignar un valor de un tipo diferente a esa variable sin ningún problema.

```julia
my_answer = "42"
```

Output:
```
"42"
```

```julia
typeof(my_answer)
```

Output:
```
String
```

## Cómo comentar

```julia
# Se pueden dejar comentarios de una sola línea usando el signo de "gato"
```

```julia
#=

Para comentarios multilínea,
se utiliza la secuencia '#= =#'.

=#
```

## Sintaxis para matemáticas básicas

```julia
sum = 3 + 7
```

Output:
```
10
```

```julia
difference = 10 - 3
```

Output:
```
7
```

```julia
product = 20 * 5
```

Output:
```
100
```

```julia
quotient = 100 / 10
```

Output:
```
10.0
```

```julia
power = 10 ^ 2
```

Output:
```
100
```

```julia
modulus = 101 % 2
```

Output:
```
1
```

### Ejercicios

#### 1.1 
Busca la documentación para la función `convert`.

```julia

```

#### 1.2 
Asigna `365` a una variable llamada `days`. Convierte `days` a un flotante y asígnala a la variable `days_float`.

```julia
# escribe tu respuesta
days = 365
days_float = convert(Float64, days)
```

Output:
```
365.0
```

<!-- grade_id: cell-715f78016beb0489, points: 1, grade: true, solution: false -->
```julia
@assert days == 365
@assert days_float == 365.0

```

#### 1.3 
Observa qué sucede cuando se ejecuta lo siguiente:

```julia
convert(Int64, "1")
```
and

```julia
parse(Int64, "1")
```

```julia
convert(Int64, "1")
```

Output:
```
Stacktrace:
 [1] top-level scope at In[24]:1
 [2] include_string(::Function, ::Module, ::String, ::String) at .\loading.jl:1091
```

```julia
parse(Int64, "1")
```

Output:
```
1
```

```julia

```