# Despacho múltiple (*multiple dispatch*)

En este *notebook*, exploramos el despacho múltiple (*multiple dispatch*), que es una característica clave de Julia.

El despacho múltiple es hace que el *software* sea **genérico** y **rápido**.

#### Comenzando con lo conocido

Para entender el despacho múltiple en Julia, comencemos con lo que ya hemos visto.

Podemos declarar funciones en Julia sin darle mucha información sobre los tipos de los argumentos de entrada que la función recibirá:

```julia
f(x) = x^2
```

Output:
```
f (generic function with 1 method)
```

y después Julia determinará por su cuenta qué tipos de argumentos de entrada tienen sentido y cuáles no:

```julia
f(10)
```

Output:
```
100
```

```julia
f([1, 2, 3])
```

Output:
```
[3] f(::Array{Int64,1}) at .\In[1]:1
 [4] top-level scope at In[3]:1
 [5] include_string(::Function, ::Module, ::String, ::String) at .\loading.jl:1091
```

#### Especificando los tipos de nuestros argumentos de entrada

Sin embargo, también tenemos la *opción* de decirle a Julia de forma explícita qué tipos están permitidos en nuestros argumentos de entrada.

Por ejemplo, escribamos una función `foo` que solo tome cadenas como entradas.

```julia
foo(x::String, y::String) = println("¡Mis entradas x y y son ambas cadenas!")
```

Output:
```
foo (generic function with 1 method)
```

Aquí vemos que para poder restringir el tipo de `x` y `y` a `String`, solo debemos escribir después del nombre del argumento de entrada dos dos puntos (::) y la palabra reservada `String`.

Ahora, veremos que `foo` funciona con cadenas y no funciona con otros tipos de argumentos de entrada.

```julia
foo("Hola", "Hi")
```

Output:
```
¡Mis entradas x y y son ambas cadenas!
```

```julia
foo(3, 4)
```

Output:
```
Stacktrace:
 [1] top-level scope at In[6]:1
 [2] include_string(::Function, ::Module, ::String, ::String) at .\loading.jl:1091
```

Para hacer que `foo` trabaje en entradas enteras (de tipo `Int`), solo ponemos `::Int` en nuestros argumentos de entrada cuando declaramos `foo`.

```julia
foo(x::Int, y::Int) = println("Mis entradas x y y son ambas enteras")
```

Output:
```
foo (generic function with 2 methods)
```

```julia
foo(3, 4)
```

Output:
```
Mis entradas x y y son ambas enteras
```

Ahora, `foo` funciona con enteros. Pero mira que `foo` también aún funciona cuando `x` y `y` son cadenas.

```julia
foo("hello", "hi!")
```

Output:
```
¡Mis entradas x y y son ambas cadenas!
```

Esto comienza a llegar al corazón del despacho múltiple. Cuando declaramos

```julia
foo(x::Int, y::Int) = println("Mis entradas x y y son ambas enteras")
```

no sobreescribimos o reemplazamos

```julia
foo(x::String, y::String)
```

En su lugar, solo agregamos un ***método*** adicional a la ***función genérica*** llamada `foo`.

Una ***función genérica*** es el concepto abstracto asociado con una operación en particular.

Por ejemplo, la función genérica `+` representa el concepto de adición.

Un ***método*** es una implementación específica de una función genérica para *tipos de argumentos particulares*.

Por ejemplo, `+` tiene métodos que aceptan números de punto flotante, enteros, matrices, etc.

Podemos usar los `métodos` para ver cuántos métodos existen para `foo`.

```julia
methods(foo)
```

Output:
```
# 2 methods for generic function "foo":
[1] foo(x::Int64, y::Int64) in Main at In[7]:1
[2] foo(x::String, y::String) in Main at In[4]:1
```

Nota aparte: ¿cuántos métodos piensas que hay para la suma?

```julia
+(2, 3, 5)
```

Output:
```
10
```

```julia
methods(+)
```

Output:
```
# 184 methods for generic function "+":
[1] +(x::Float32, y::Float32) in Base at float.jl:400
[2] +(a::Pkg.Resolve.VersionWeight, b::Pkg.Resolve.VersionWeight) in Pkg.Resolve at D:\Programas\Julia 1.5.3\share\julia\stdlib\v1.5\Pkg\src\Resolve\versionweights.jl:22
[3] +(dt::Dates.DateTime, z::Dates.Month) in Dates at D:\Programas\Julia 1.5.3\share\julia\stdlib\v1.5\Dates\src\arithmetic.jl:47
[4] +(dt::Dates.DateTime, y::Dates.Year) in Dates at D:\Programas\Julia 1.5.3\share\julia\stdlib\v1.5\Dates\src\arithmetic.jl:23
[5] +(x::Dates.DateTime, y::Dates.Period) in Dates at D:\Programas\Julia 1.5.3\share\julia\stdlib\v1.5\Dates\src\arithmetic.jl:76
[6] +(a::Float16, b::Float16) in Base at float.jl:398
[7] +(level::Base.CoreLogging.LogLevel, inc::Integer) in Base.CoreLogging at logging.jl:116
[8] +(x::Dates.Date, y::Dates.Day) in Dates at D:\Programas\Julia 1.5.3\share\julia\stdlib\v1.5\Dates\src\arithmetic.jl:74
[9] +(x::Dates.Date, y::Dates.Week) in Dates at D:\Programas\Julia 1.5.3\share
```

Así que ahora podemos llamar `foo` con enteros o cadenas. Cuando llamas `foo` para un conjunto específico de argumentos, Julia inferirá los tipos de las entradas y *despachará* el método apropiado. *Esto* es el despacho múltiple.

El despacho múltiple hace que nuestro código sea genérico y rápido. Nuestro código puede ser genérico y flexible porque podemos escribir código en términos de operaciones abstractas tales como la suma y la multiplicación, en lugar de en términos de implementaciones específicas. Al mismo tiempo, nuestro código corre de forma rápida porque Julia es capaz de llamar métodos eficientes para los tipos relevantes.

Para ver qué métodos es el que está siendo despachado cuando llamamos una función genérica, podemos usar la macro `@which`:

```julia
@which foo(3, 4)
```

Output:
```
foo(x::Int64, y::Int64) in Main at In[7]:1
```

Veamos qué sucede cuando usamos `@which` con el operador de suma.

```julia
@which 3 + 3.0
```

Output:
```
+(x::Number, y::Number) in Base at promotion.jl:311
```

Y podemos continuar agregando otros métodos a nuestra función genérica `foo`. Agreguemos uno que tome el ***tipo abstracto*** `Number`, que incluye los subtipos tales como `Int`, `Float64` y otros objetos pudiera pensar como números:

```julia
foo(x::Number, y::Number) = println("Mis entradas x y y son ambas números")
```

Output:
```
foo (generic function with 3 methods)
```

Este método para `foo` funcionará en, por ejemplo, números de punto flotante:

```julia
foo(3.0, 4.0)
```

Output:
```
Mis entradas x y y son ambas números
```

También podemos agregar un método que tome entradas de cualquier tipo:

```julia
foo(x, y) = println("Acepto entradas de cualquier tipo")
```

Output:
```
foo (generic function with 4 methods)
```

Con los métodos que hemos escrito hasta ahora para `foo`, este método será llamado cuando le pasemos cosas que no son números a `foo`:

```julia
v = rand(3)
foo(v, v)
```

Output:
```
Acepto entradas de cualquier tipo
```

```julia
@which foo(v, v)
```

Output:
```
foo(x, y) in Main at In[21]:1
```

### Ejercicios

#### 7.1

Extiende la función `foo` al agregar un método que tome únicamente un argumento de entrada que sea de tipo `Bool` y regrese "foo con un booleano"

```julia
function foo(x::Bool)::String
  return "foo con un booleano"
end
```

Output:
```
foo (generic function with 5 methods)
```

```julia
methods(foo)
```

Output:
```
# 5 methods for generic function "foo":
[1] foo(x::Bool) in Main at In[24]:1
[2] foo(x::Int64, y::Int64) in Main at In[7]:1
[3] foo(x::String, y::String) in Main at In[4]:1
[4] foo(x::Number, y::Number) in Main at In[19]:1
[5] foo(x, y) in Main at In[21]:1
```

#### 7.2

Verifica que el método despachado al ejecutar 
```julia
foo(true)
```
es el que escribiste.

```julia
@which foo(true)
```

Output:
```
foo(x::Bool) in Main at In[24]:1
```

<!-- grade_id: cell-14072e60ae07c1a2, points: 1, grade: true, solution: false -->
```julia
@assert foo(true) == "foo con un booleano"
```

```julia

```